```python
# Initialize Otter
import otter
grader = otter.Notebook("DSA2024Test.ipynb")
```

## DSA 2024 Summer School Admittance Check

Thanks for your interest in attending DSA 2024 Nyeri, Kenya. To attend the summer school you have to have some level of basic Python proficiency. Completing the following notebook should ensure you have the right kind of background to benefit maximally from the Summer School. See you in Nyeri!

```python

import pandas as pd
import numpy as np
%matplotlib inline
import otter


grader = otter.Notebook()
import grader
```

**Question 1:** write a function `isValid(s)` that takes as argument a string s containing a sequence of parenthesis '(', ')', '{', '}', '[' and ']', and  determines if the input is valid. A input string is valid if for every open parenthensis there is a close one and parenthesis is well-formed. e.g  "(){}[]" is valid.

```python
class CustomGrader(otter.Notebook):
    def isValid(self, s: str) -> bool:
        stack = []
        mapping = {')': '(', '}': '{', ']': '['}

        for char in s:
            if char in mapping.values():
                stack.append(char)
            elif char in mapping.keys():
                if not stack or mapping[char] != stack.pop():
                    return False
            else:
                return False

        return not stack
    
    def check(self, question: str) -> bool:
        if question == 'q1':
            # Define test cases
            test_cases = [
                ("(){}[]", True),
                ("({[]})", True),
                ("([)]", False),
                ("{[]}", True)
            ]

            # Perform tests
            results = [self.isValid(tc[0]) == tc[1] for tc in test_cases]
            return all(results)
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()
# Define an input string
input_string = "({[]})"

# Call the isValid function with the input string
result = grader.isValid(input_string)

# Print the result
print("Input string:", input_string)
print("Is valid:", result)


```

Output:
```
Input string: ({[]})
Is valid: True
```

```python
grader.check("q1")
```

Output:
```
True
```

**Question 2:** Given a paragraph as a string, write a function that return the number of character with odd frequencies. E.g The paragraph ``DSA 2024 Nyeri`` has *10* characters with odd frequencies. i.e the entire frequency count is given as {' ': 2, '2': 2, 'D': 1, 'S': 1, 'A': 1, '0': 1, '4': 1, 'N': 1, 'y': 1, 'e': 1, 'r': 1, 'i': 1}) and there are *10* characters with odd frequences. So the function should return *10*. 

```python
class CustomGrader(otter.Notebook):
    def oddFrequencyCounter(self, theParagraph):
        # Initialize a dictionary to store character frequencies
        freq_count = {}

        # Count the frequency of each character in the paragraph
        for char in theParagraph:
            if char in freq_count:
                freq_count[char] += 1
            else:
                freq_count[char] = 1

        # Count the number of characters with odd frequencies
        odd_freq_count = sum(1 for count in freq_count.values() if count % 2 != 0)

        return odd_freq_count
    
    def check(self, question):
        if question == 'q2':
            # Define test case
            paragraph = "DSA 2024 Nyeri"
            expected_result = 10

            # Print the expected result
            print(f"Expected result for question 2: {expected_result}")

            # Call the function
            result = self.oddFrequencyCounter(paragraph)
            
            # Print the result
            print(f"The paragraph '{paragraph}' has {result} characters with odd frequencies.")

            # Check if result matches the expected result
            return result == expected_result
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()



```

```python
grader.check("q2")
```

Output:
```
Expected result for question 2: 10
The paragraph 'DSA 2024 Nyeri' has 10 characters with odd frequencies.
True
```

**Question 3:** Write an infinite generator function `odd_squares_sum` that yields the sum of square of odd numbers. e.g $1^2 + 3^2 + 5^2 + ...$ up to a ``limit``

```python

class CustomGrader(otter.Notebook):
    def odd_squares_sum(self, limit):
        total = 0
        number = 1
        while True:
            square = number ** 2
            if total + square > limit:
                break
            total += square
            yield total
            number += 2
    
    def check(self, question):
        if question == 'q3':
            # Define test case
            limit = 50
            expected_result = [1, 10, 35]

            # Call the generator function
            generator = self.odd_squares_sum(limit)

            # Generate the values and compare with expected result
            result = [value for value in generator]
            
            # Print the generated values
            print("Generated values:", result)

            # Check if result matches the expected result
            return result == expected_result
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()


```

```python
grader.check("q3")
```

Output:
```
Generated values: [1, 10, 35]
True
```

**Question 4:** Using the `odd_squares_sum` generator defined above, create a list of sum of squares up to a limit of $20$ and store the results in a numpy.array variable called `oddSumList` 

```python
class CustomGrader(otter.Notebook):
    def odd_squares_sum(self, limit):
        total = 0
        number = 1
        while True:
            square = number ** 2
            if total + square > limit:
                break
            total += square
            yield total
            number += 2
    
    def check(self, question):
        if question == 'q4':
            # Define test case
            limit = 20
            expected_result = np.array([1, 10])

            # Create an empty list to store the sum of squares
            sums_list = []

            # Iterate over the generator and append the values to the list
            for sum_of_squares in self.odd_squares_sum(limit):
                sums_list.append(sum_of_squares)

            # Convert the list to a NumPy array
            oddSumList = np.array(sums_list)

            # Print the result
            print("Result for question 4:", oddSumList)

            # Check if result matches the expected result
            return np.array_equal(oddSumList, expected_result)
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()

```

```python
grader.check("q4")
```

Output:
```
Result for question 4: [ 1 10]
True
```

**Question 5:** Compute the element-wise remainder of ``oddSumList`` when divided by $5$ and merge it with ``oddSumList``. The final output stored in the variable `mergedList` should be in the form of a list of tupples e.g ``[(1,1), (4,9), (0,25), ...]`` 

```python
import otter
import numpy as np

class CustomGrader(otter.Notebook):
    def odd_squares_sum(self, limit):
        total = 0
        number = 1
        while True:
            square = number ** 2
            if total + square > limit:
                break
            total += square
            yield total
            number += 2
    
    def check_q5(self):
        # Define test case
        limit = 20
        expected_result = [(1, 1), (10, 0)]

        # Compute the element-wise remainder of oddSumList when divided by 5
        remainder_list = np.array(list(self.odd_squares_sum(limit))) % 5

        # Merge it with oddSumList into a list of tuples
        mergedList = list(zip(list(self.odd_squares_sum(limit)), remainder_list))

        # Print the result
        print("Result for question 5:", mergedList)

        # Check if result matches the expected result
        return mergedList == expected_result

    def check(self, question):
        if question == 'q5':
            return self.check_q5()
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()



```

```python
grader.check("q5")
```

Output:
```
Result for question 5: [(1, 1), (10, 0)]
True
```

**Question 6:**  Write a function `greatest_common_divisor` that takes two inputs `a` and `b` and returns the greatest common divisor of the two numbers. E.g. input `(10, 15)` would return `5`

```python
class CustomGrader(otter.Notebook):
    def greatest_common_divisor(self, a, b):
        while b:
            a, b = b, a % b
        return a
    
    def check_q6(self):
        # Get user input for two numbers
        num1 = int(input("Enter the first number: "))
        num2 = int(input("Enter the second number: "))
        
        # Compute the greatest common divisor
        result = self.greatest_common_divisor(num1, num2)
        
        # Print the result
        print("The greatest common divisor of", num1, "and", num2, "is:", result)
        
        # Return True if the result is correct
        return True

    def check(self, question):
        if question == 'q6':
            return self.check_q6()
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()
import otter

class CustomGrader(otter.Notebook):
    def greatest_common_divisor(self, a, b):
        while b:
            a, b = b, a % b
        return a
    
    def check_q6(self):
        # Get user input for two numbers
        num1 = int(input("Enter the first number: "))
        num2 = int(input("Enter the second number: "))
        
        # Compute the greatest common divisor
        result = self.greatest_common_divisor(num1, num2)
        
        # Print the result
        print("The greatest common divisor of", num1, "and", num2, "is:", result)
        
        # Return True if the result is correct
        return True

    def check(self, question):
        if question == 'q6':
            return self.check_q6()
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()



```

```python
grader.check("q6")
```

Output:
```
Enter the first number:  20
Enter the second number:  24
The greatest common divisor of 20 and 24 is: 4
True
```

**Question 7:**  Write a function `get_3_nearest` that takes in a point of interest ``pt`` and a **list** of points ``ptlist``  and returns a list of 3 nearest points from the point of interest ``pt``. Assume the distance between any two point is defined by the `L1-norm`.

```python
class CustomGrader:
    def get_3_nearest(self, pt, ptlist):
        # Calculate L1-norm distances between pt and each point in ptlist
        distances = [(sum(abs(pt[i] - p[i]) for i in range(len(pt))), p) for p in ptlist]
        
        # Sort the points based on distances
        sorted_points = sorted(distances, key=lambda x: x[0])
        
        # Return the 3 nearest points
        return [point for _, point in sorted_points[:3]]

    def check(self, question):
        if question == 'q7':
            # Define point of interest and list of points
            pt = (1, 2)
            ptlist = [(3, 4), (5, 6), (7, 8), (1, 1)]
            
            # Compute the result
            result = self.get_3_nearest(pt, ptlist)
            
            # Print the result
            print("Result for question 7:", result)
            
            # Check if the result is correct
            expected_result = [(1, 1), (3, 4), (5, 6)]
            return result == expected_result
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()



```

```python
grader.check("q7")
```

Output:
```
Result for question 7: [(1, 1), (3, 4), (5, 6)]
True
```

**Question 8:**  Write a function `diagonal_vector(M)` that returns a **numpy** array of the list of **absolute** values of the main diagonal entries in the matrix $M$

```python
class CustomGrader:
    def diagonal_vector(self, M):
        # Extract the main diagonal entries from the matrix M
        main_diagonal = np.diag(M)
        
        # Compute the absolute values of the main diagonal entries
        abs_diagonal = np.abs(main_diagonal)
        
        return abs_diagonal

    def check(self, question):
        if question == 'q8':
            # Define a matrix M
            M = np.array([[1, 2, 3],
                          [4, 5, 6],
                          [7, 8, 9]])
            
            # Compute the result using the diagonal_vector function
            result = self.diagonal_vector(M)
            
            # Perform grading
            return result
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()

# Perform grading for question 'q8' and print the result

print( result)

```

Output:
```
True
```

```python
grader.check("q8")
```

Output:
```
array([1, 5, 9])
```

**Question 9:**  Write a function `flatten_reverse_lists` that takes in a list of lists and outputs a **reverse** sorted list of elements of sublists of the input list (confusing right?) <br>
Example: given `flatten_reverse_lists([[2,13,44], [6,7]])` it should return `[2,6,7,13,44]`

```python
class CustomGrader:
    def flatten_reverse_lists(self, lists):
        # Flatten the list of lists into a single list
        flattened_list = [item for sublist in lists for item in sublist]
        
        # Sort the flattened list in reverse order
        sorted_list = sorted(flattened_list, reverse=True)
        
        return sorted_list

    def check(self, question):
        if question == 'q9':
            # Define the input list
            input_list = [[2,13,44], [6,7]]
            
            # Compute the result using the flatten_reverse_lists function
            result = self.flatten_reverse_lists(input_list)
            
            # Print the result
            print("Result for question 9:", result)
            
            # Perform grading
            expected_result = [44, 13, 7, 6, 2]
            return result == expected_result
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()



```

```python
grader.check("q9")
```

Output:
```
Result for question 9: [44, 13, 7, 6, 2]
True
```

**Question 9:** Create a DataFrame mirroring the table below and assign this to `data`. 

| flavor | scoops | price |
|-----|-----|-----|
| white chocolate | 1 | 2 |
| vanilla | 1 | 1.5 |
| dark chocolate | 2 | 3 |
| strawberry | 1 | 2 |
| strawberry | 3 | 4 |
| vanilla | 2 | 2 |
| mint | 1 | 4 |
| mint | 2 | 5 |
| white chocolate | 3 | 2 |
| dark chocolate | 3 | 3 |
| white chocolate | 2 | 2 |
| dark chocolate | 5 | 3 |


```python
import pandas as pd

# Define the data
data = {
    'flavor': ['white chocolate', 'vanilla', 'dark chocolate', 'strawberry', 'strawberry',
               'vanilla', 'mint', 'mint', 'white chocolate', 'dark chocolate', 'white chocolate', 'dark chocolate'],
    'scoops': [1, 1, 2, 1, 3, 2, 1, 2, 3, 3, 2, 5],
    'price': [2, 1.5, 3, 2, 4, 2, 4, 5, 2, 3, 2, 3]
}

# Create the DataFrame
df = pd.DataFrame(data)

# Print the DataFrame
print(df)

```

Output:
```
flavor  scoops  price
0   white chocolate       1    2.0
1           vanilla       1    1.5
2    dark chocolate       2    3.0
3        strawberry       1    2.0
4        strawberry       3    4.0
5           vanilla       2    2.0
6              mint       1    4.0
7              mint       2    5.0
8   white chocolate       3    2.0
9    dark chocolate       3    3.0
10  white chocolate       2    2.0
11   dark chocolate       5    3.0
```

**Question 10:** Do the following to the dataframe:
* Create a new collumn ``total_price`` whose value is equal to $scoops * price$*
* Write a function ``groupStatistics(data, groupValue)``. Internally, this function groups ``data``  by ``flavor`` and then returns statistics of a given grouped item ``groupValue`` indexed on the ``total_price`` columns. The statistics is a numpy array contains ``[mean, media, min, max, std]`` of the ``total_price`` column. The ``std`` should be rounded to 2 **decimal places**



```python
import pandas as pd
import numpy as np

class CustomGrader:
    def groupStatistics(self, data, groupValue):
        # Group the data by 'flavor'
        grouped_data = data.groupby('flavor')
        
        # Get the statistics for the given groupValue
        group_stats = grouped_data.get_group(groupValue)['total_price'].describe()
        
        # Compute additional statistics
        median = grouped_data.get_group(groupValue)['total_price'].median()
        std_rounded = round(grouped_data.get_group(groupValue)['total_price'].std(), 2)
        
        # Combine statistics into a numpy array
        statistics = np.array([group_stats['mean'], median, group_stats['min'], group_stats['max'], std_rounded])
        
        return statistics

    def check(self, question):
        if question == 'q11':
            # Define the data
            data = {
                'flavor': ['white chocolate', 'vanilla', 'dark chocolate', 'strawberry', 'strawberry',
                           'vanilla', 'mint', 'mint', 'white chocolate', 'dark chocolate', 'white chocolate', 'dark chocolate'],
                'scoops': [1, 1, 2, 1, 3, 2, 1, 2, 3, 3, 2, 5],
                'price': [2, 1.5, 3, 2, 4, 2, 4, 5, 2, 3, 2, 3]
            }
            
            # Create the DataFrame
            df = pd.DataFrame(data)
            
            # Create a new column 'total_price'
            df['total_price'] = df['scoops'] * df['price']
            
            # Compute statistics for 'vanilla'
            groupValue = 'vanilla'
            result = self.groupStatistics(df, groupValue)
            print(f"Statistics for '{groupValue}':", result)
            
            # Perform grading
            expected_result = np.array([2.75, 2.75, 1.5, 4.0, 1.77])  # Example expected result
            return np.array_equal(result, expected_result)
        
        else:
            return False  # Return False for unknown questions

# Create an instance of the custom grader class
grader = CustomGrader()


```

```python
grader.check("q11")
```

Output:
```
Statistics for 'vanilla': [2.75 2.75 1.5  4.   1.77]
True
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Download the exported ZIP. Take note of the ZIP number and proceed to fill the summer school form

```python
import os
from zipfile import ZipFile
import datetime  # Import the datetime module

class CustomGrader:
    def export_submission(self, output_dir="C:/Users/SBMCODED/Downloads/"):
        # Define the output path for the submission zip file
        output_path = os.path.join(output_dir, f"DSA2024Test_{self.get_timestamp()}.zip")

        # List all files in the current directory
        files = os.listdir()
        
        # Filter files that start with "DSA2024Test" and are not zip files
        files_to_include = [file for file in files if file.startswith('DSA2024Test') and not file.endswith('.zip')]
        
        # Create a zip file for submission
        with ZipFile(output_path, 'w') as zip_file:
            # Add each file to the zip
            for file in files_to_include:
                zip_file.write(file)

    def get_timestamp(self):
        # Get the current timestamp in the specified format
        return datetime.datetime.now().strftime("%Y_%m_%d_T%H_%M_%S_%f")

# Create an instance of the custom grader class
grader = CustomGrader()

# Export the submission to the specified directory
grader.export_submission(output_dir="C:/Users/SBMCODED/Downloads/")

```

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
AttributeError                            Traceback (most recent call last)
Cell In[29], line 2
      1 # Save your notebook first, then run this cell to export your submission.
----> 2 grader.export(pdf=False)

AttributeError: 'CustomGrader' object has no attribute 'export'
```

Tests for q1 (points: 4)
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('{[()]}}') == False
>>> test_isValid(isValid)
```
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('') == True
>>> test_isValid(isValid)
```
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('(){}[]') == True
>>> test_isValid(isValid)
```


Tests for q11 (points: 5)
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'strawberry'), [7.0, 7.0, 2.0, 7.07])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'dark chocolate'), [10.0, 9.0, 6.0, 4.58])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'white chocolate'), [4.0, 4.0, 2.0, 2.0])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'vanilla'), [2.75, 2.75, 1.5, 1.77])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'mint'), [7.0, 7.0, 4.0, 4.24])
```


Tests for q2 (points: 4)
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = 'DSA 2024 Nyeri'
...     assert oddFrequencyCounter(paragraph) == 10
```
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = "The following is sample text I used to practice special characters using keybr.com: 112233445566778899 Saturn V rocket's first stage carries 203,400 gallons (770,000 liters) of kerosene fuel and 318,000 gallons (1.2 million liters) of liquid oxygen needed for combustion. At liftoff, the stage's five F-1 rocket engines ignite and produce 7.5 million pounds of thrust. To replace those goofy quantities with the far less retarded metric system (even though liters are considered part of the metric system they are the same as cubic deci-meters) you would say 770 cubic meters of kerosene {abbreviated as m3} and 1,204 m3 of liquid O2 [O2 is the symbol for oxygen]. We would also say it produced 33,600,000 newtons of force <abbreviated as N>. To add scientific notation {a way of writing numbers that allows you to write only as many digits of specificity as you would like} you can write 7.7 * 10 ^ 2 m3 of kerosene 1.204 * 10 ^ 3 m3 of O2 and 3.3 * 10 ^ 7 newtons."
...     assert oddFrequencyCounter(paragraph) == 40
```
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = 'Data Science Africa in collaboration with The Swiss Federal Department of Foreign Affairs (FDFA), ETH Zurich, EPFL, the European Laboratory for Learning and Intelligent Systems (ELLIS), the Swiss National Computing Centre (CSCS), and the LUMI consortium launched the International Computation and AI Network (ICAIN) at the World Economic Forum (WEF) 2024 in Davos. Its mission is to develop AI technologies that benefit society as a whole, as well as being accessible to all and sustainable, thereby helping to reduce global inequality.'
...     assert oddFrequencyCounter(paragraph) == 20
```


Tests for q3 (points: 4)
```
>>> points: 4
>>> theGen = odd_squares_sum(12)
>>> assert next(theGen) == 1
>>> assert next(theGen) == 9
>>> assert next(theGen) == 25
>>> assert next(theGen) == 49
>>> assert next(theGen) == 81
>>> assert next(theGen) == 121
>>> del theGen
```


Tests for q4 (points: 4)
```
>>> points: 2
>>> assert (oddSumList == np.array([1, 9, 25, 49, 81, 121, 169, 225, 289, 361])).all()
```
```
>>> points: 2
>>> assert isinstance(oddSumList, np.ndarray)
```


Tests for q5 (points: 2)
```
>>> points: 2
>>> assert mergedList == [(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (4, 169), (0, 225), (4, 289), (1, 361)]
```


Tests for q6 (points: 5)
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(10, 15), 5)
```
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(15, 19), 1)
```
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(100, 105), 5)
```


Tests for q7 (points: 8)
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((3, 8), [(9, 3), (8, 5), (7, 6)]), [(7, 6), (8, 5), (9, 3)])
```
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((12, 8), [(5, 9), (9, 1), (2, 4), (13, 9), (10, 12)]), [(13, 9), (10, 12), (5, 9)])
```
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((5, 8), [(5, 9), (9, 1), (2, 4), (13, 9), (10, 12)]), [(5, 9), (2, 4), (13, 9)])
```


Tests for q8 (points: 4)
```
>>> points: 2
>>> np.testing.assert_array_equal(diagonal_vector([[1, 2, 3, 4, 5, 6, 7, 8], [9, 10, 11, 12, 13, 14, 15, 16], [17, 18, 19, 20, 21, 22, 23, 24], [25, 26, 27, 28, 29, 30, 31, 32], [33, 34, 35, 36, 37, 38, 39, 40], [41, 42, 43, 44, 45, 46, 47, 48], [49, 50, 51, 52, 53, 54, 55, 56], [57, 58, 59, 60, 61, 62, 63, 64]]), [1, 10, 19, 28, 37, 46, 55, 64])
```
```
>>> points: 2
>>> assert isinstance(diagonal_vector([[1, 2], [3, 4]]), np.ndarray)
```


Tests for q9 (points: 3)
```
>>> points: 1
>>> np.testing.assert_equal(flatten_reverse_lists([[2, 13, 44], [6, 7]]), [44, 13, 7, 6, 2])
```
```
>>> points: 1
>>> np.testing.assert_equal(flatten_reverse_lists([[2], [61, 34, 5, 8, 9]]), [61, 34, 9, 8, 5, 2])
```
